rule_files:
  - ../slo_burn_rules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'slo:read_latency_errors:burnrate5m{service="graphql"}'
        values: '0 0 0 15 15 15 0 0'
    alert_rule_test:
      - alertname: ReadLatencyBudgetBurnFast
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
  - interval: 1m
    input_series:
      - series: 'slo:read_latency_errors:burnrate5m{service="graphql"}'
        values: '0 0 0 5 5 5 0 0'
    alert_rule_test:
      - alertname: ReadLatencyBudgetBurnFast
        eval_time: 6m
        exp_alerts: []

